<a href="https://colab.research.google.com/github/tmckim/materials-fa24-colab/blob/main/review/Q3_ReviewNotebook.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>


## Before you start - Save this notebook!

When you open a new Colab notebook from the WebCampus (like you hopefully did for this one), you cannot save changes. So it's  best to store the Colab notebook in your personal drive `"File > Save a copy in drive..."` **before** you do anything else.

The file will open in a new tab in your web browser, and it is automatically named something like: "**Copy of Q3_ReviewNotebook.ipynb**". You can rename this to just the title of the assignment "**Q3_ReviewNotebook.ipynb**". Make sure you do keep an informative name (like the name of the assignment) so that you know which files to submit back to WebCampus for grading! More instructions on this are at the end of the notebook.


**Where does the notebook get saved in Google Drive?**

By default, the notebook will be copied to a folder called “Colab Notebooks” at the root (home directory) of your Google Drive. If you use this for other courses or personal code notebooks, I recommend creating a folder for this course and then moving the assignments AFTER you have completed them. <br>

I also recommend you give the folder where you save your notebooks^ a different name than the folder we create below that will store the notebook resources you need each time you work through a course notebook. This includes any data files you will need, links to the images that appear in the notebook, and the files associated with the autograder for answer checking.<br>
You should select a name other than '**NS499-DataSci-course-materials**'. <br>
This folder gets overwritten with each assignment you work on in the course, so you should **NOT** store your notebooks in this folder that we use for course materials! <br><br>For example, you could create a folder called 'NS499-**notebooks**' or something along those lines.
___

# Import & Setup

```python
# Step 1
# Setup and add files needed to access gdrive
from google.colab import drive                                   # these lines mount your gdrive to access the files we import below
drive.mount('/content/gdrive', force_remount=True)
```

```python
# Step 2
# Change directory to the correct location in gdrive (modified way to do this from before)
import os
os.chdir('/content/gdrive/MyDrive/NS499-DataSci-course-materials/')
```

```python
# Step 3
# Remove the files that were previously there- we will replace with all the old + new ones for this assignment
!rm -r materials-fa24-colab
```

```python
# Step 4
# These lines clone (copy) all the files you will need from where I store the code+data for the course (github)
# Second part of the code copies the files to this location and folder in your own gdrive
!git clone https://github.com/tmckim/materials-fa24-colab '/content/gdrive/My Drive/NS499-DataSci-course-materials/materials-fa24-colab/'
```

```python
# Step 5
# Change directory into the folder where the resources for this assignment are stored in gdrive (modified way from before)
os.chdir('/content/gdrive/MyDrive/NS499-DataSci-course-materials/materials-fa24-colab/review/')
```

```python
# Import packages and other things needed
# Don't change this cell; Just run this cell
# If you restart colab, make sure to run this cell again after the first ones above^

# These lines import the Numpy and Datascience modules.
import numpy as np
from datascience import *

# These lines do some fancy plotting magic.
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', FutureWarning)
```

## Python Expressions & Methods

```python
# Q1
sum(make_array(True,False,False,True,False))
```

```python
# Q2
a = make_array(8,1,22,14,3,10,4,18)
np.count_nonzero(a>10)
```

```python
# Q3
first = np.arange(4)
np.append(first, make_array(6,10))
```

```python
# Q4
x = 7
y = 3
10 > x < 5
```

```python
#@title Animals table
# Make the animals table in datascience package
animal = make_array("dog", "cat", "dog", "dog", "hamster", "cat")
color = make_array("brown", "black", "white", "brown", "orange", "black")


animals = Table().with_columns('animal', animal, 'color', color)

animals
```

Output:
```
animal  | color
dog     | brown
cat     | black
dog     | white
dog     | brown
hamster | orange
cat     | black
```

```python
# Q5
animals.pivot('animal', 'color')
```

```python
# Q6
animals.group('color')
```

## Simulations & Defining Functions with Roulette

A Nevada roulette wheel has 38 pockets and a small ball that rests on the wheel. When the wheel is spun, the ball comes to rest in one of the 38 pockets. That pocket is declared the winner.

The pockets are labeled 0, 00, 1, 2, 3, 4, ... , 36. Pockets 0 and 00 are green, and the other pockets are alternately red and black. The table `wheel` is a representation of a Nevada roulette wheel. **Note that *both* columns consist of strings.** Below is an example of a roulette wheel!

<img src="roulette_wheel.jpeg" width="330px">

Run the cell below to load the `wheel` table.

```python
wheel = Table.read_table('roulette_wheel.csv')
wheel
```

Output:
```
Pocket | Color
0      | green
0      | green
1      | red
2      | black
3      | red
4      | black
5      | red
6      | black
7      | red
8      | black
... (28 rows omitted)
```

### Betting on Red ###
If you bet on *red*, you are betting that the winning pocket will be red. This bet *pays 1 to 1*. That means if you place a one-dollar bet on red, then:

- If the winning pocket is red, you gain 1 dollar. That is, you get your original dollar back, plus one more dollar.
- if the winning pocket is not red, you lose your dollar. In other words, you gain -1 dollars.

**Hint:** This is like what we practiced in class when we compared my roll of a die to yours with simulation. One outcome was positive and one was negative.

**Question 1.** Define a function `dollar_bet_on_red` that takes the name of a color and returns your gain in dollars if that color had won and you had placed a one-dollar bet on red. Remember that the gain can be negative. Make sure your function returns an integer.

*Note:* You can assume that the only colors that will be passed as arguments are red, black, and green. Your function doesn't have to check that.


```python
def dollar_bet_on_red(color):
```

Run the cell below to make sure your function is working.

```python
print(dollar_bet_on_red('green'))
print(dollar_bet_on_red('black'))
print(dollar_bet_on_red('red'))
```

**Question 2.** Now it's time to apply the function you just wrote. Your code should use the function `dollar_bet_on_red` on the correct column from the `wheel` table. The output array `red_winnings` should contain the gain in dollars if your pocket won and ou had bet one dollar on red.



```python
red_winnings = ...
```

**Question 3:** Your friend Steve was trying to be helpful (and he doesn't know how to code!), so he compiled the data into a table called `winnings`. You would like to add this data to the original `wheel` table. Create a new table called `all_data` that contains 3 columns: `Pocket`, `Color`, and `Red Wins` by using a method that allows you to combine the two tables.

```python
#@title updated wheel table
wheel_updated = wheel.with_columns('Red Wins', red_winnings)
wheel_updated
```

```python
#@title Winnings table
winnings = wheel_updated.select('Pocket', 'Red Wins').relabel('Pocket', 'Number')
winnings
```

Output:
```
Number | Red Wins
0      | -1
0      | -1
1      | 1
2      | -1
3      | 1
4      | -1
5      | 1
6      | -1
7      | 1
8      | -1
... (28 rows omitted)
```

**Question 4:** Below we have defined a function that simulates one round of a spin on the roulette wheel. We use this function and a for loop to simulate 10000 spins and store the outcomes. The outcomes (1 or -1) are then put into the `results` table. We are interested in the count of wins (1) and losses (-1) from the `results` table. Write code using a method we've learned to output a table that shows the cunt for the values in the `My winnings` column.

```python
def simulate_one_round():
  color = np.random.choice(wheel.column('Color'))
  return dollar_bet_on_red(color)
```

```python
# simulation
outcomes = make_array()
for i in np.arange(10000):
  outcomes = np.append(outcomes, simulate_one_round())

results = Table().with_column('My winnings', outcomes)
results
```

Output:
```
My winnings
-1
1
1
-1
1
1
-1
1
-1
1
... (9990 rows omitted)
```

**Question 5:** We took a random sample of 20 rows from the original `wheel` table and put those results into a new table called `twenty_bets`. However, instead of putting in the original numbers, 1 and 0 have replaced 1 and -1, respectively. You'd like to know how many times you won on your bets by summing the values in the `Wins Adjusted` column. Write code to calculate the sum from this data in the `twenty_bets` table.

```python
#@title Hidden
twenty_bets = wheel_updated.sample(20)
twenty_bets

values = make_array()

for i in np.arange(len(twenty_bets.column('Red Wins'))):
  if twenty_bets.column('Red Wins').item(i) == 1:
       values = np.append(values,True)
  else:
       values = np.append(values,False)

twenty_bets = twenty_bets.with_column('Wins Adjusted', values).drop('Red Wins')
twenty_bets
```

Output:
```
Pocket | Color | Wins Adjusted
32     | red   | 1
7      | red   | 1
36     | red   | 1
0      | green | 0
29     | black | 0
21     | red   | 1
33     | black | 0
12     | red   | 1
5      | red   | 1
19     | red   | 1
... (10 rows omitted)
```

## Conditionals, Iteration, and Randomization

**Question 1:** A research lab studies brain development and invites children, teens, or adults to participate in their studies. During each session, they collect information about each participant’s age. They want help writing code that can be used to take the participant age and convert it to a grouping variable. To do so, they ask you to define a function with the following information. <br><br>
The `age_group` function should take in the age of the participant (from 8 to 40) and assign text to indicate the corresponding group. The final output should be the text for that age group. You do not need to get that text from the group column in the table. Instead, the function should explicitly return the text as a string, depending on the age input to the function.<br><br>
Hint: Notice in the column for age, it provides ranges with brackets [ and parentheses ). Remember that this notation means that the number on the left is included in that range, but the one on the right is not. So bracket [ means inclusive and parenthesis ) means exclusive (except the last one, which includes both ends). This is important for determining what to use in your `if/elif` statements below.<br><br>


   age    |  group
----------|----------
[8,11)    | Child
[11,18)   | Teen
[18,40)   | Adult


```python
# Q1
def age_group(age):
  """ Takes age (int) and returns the corresponding group variable (str)"""
  if    :
    return
  elif :
    return
  else:
    return


```

```python
# Q2
for ... in make_array('fall', 'winter', 'spring', 'summer'):
...
```

```python
# Q3
x = 5
for i in np.arange(2,5):
  x = x + i
  print(x)
```

```python
#@title Hidden
# loop through 2,3,4
# 1st iteration: i = 2; x = 5 + 2 = 7
# 2nd iteration: i = 3; x = 7 + 3 = 10
# 3rd iteration: i = 4; x = 10 + 4 = 14
```

**Question 4:** You have an array of Halloween `candy` below. Write code that will randomly choose 2 candy options that you should buy from the store to give out for trick or treat.

```python
# Q4
candy = make_array("Reese's", "KitKat", "Sour Patch Kids", 'Twizzlers')
candy
```

```python
yummy = ...
```

# Extra Questions

### Simulating 10 Bets on Red
Roulette wheels are set up so that each time they are spun, the winning pocket is equally likely to be any of the 38 pockets regardless of the results of all other spins. Let's see what would happen if we decided to bet one dollar on red each round.

**Question 3.** Write a line of code that creates a table `ten_bets` by sampling the table `wheel` to simulate 10 spins of the roulette wheel. Your table should have the same three column labels as in `wheel`.


```python
ten_bets = ...
```

```python
# let's find the sum of our winnings
total = sum(ten_bets.column('Red Wins'))
total
```

Run the cells above a few times to see how much money you would make if you made 10 one-dollar bets on red. Making a negative amount of money doesn't feel good, but it is a reality in gambling. Casinos are a business, and they make money when gamblers lose.

**Question 4.** Let's see what would happen if you made more bets. Define a function `net_gain_red` that takes the number of bets and returns the net gain in that number of one-dollar bets on red.

*Hint 1:* You should use your `wheel` table within your function, with similar code to what you wrote in the above question. <br>
*Hint 2:* you should return the sum of the column `Red Wins` from your new table created in the function below


```python
def net_gain_red(...):
  all_bets = ...
  return ...

```

Run the cell below a few times to make sure that the results are similar to those you observed in the previous exercise.

```python
net_gain_red(10)
```

**Question 5.** Complete the cell below to simulate the net gain in 200 one-dollar bets on red, repeating the process 10,000 times. After the cell is run, `all_gains_red` should be an array with 10,000 entries, each of which is the net gain in 200 one-dollar bets on red.

*Hint:* For this question, you can assume that `numpy` has been imported as `np`.


```python
num_bets = 200
repetitions = 10000
all_gains_red = make_array()


for ... in ...(...):
    new_net_gain = ...(...)
    all_gains_red = ...(..., ...)

len(all_gains_red)
```

Run the cell below to visualize the results of your simulation.

```python
gains = Table().with_columns('Net Gain on Red', all_gains_red)
gains.hist(bins = np.arange(-80, 41, 4))
```

**Extra Thought Question:** Using the histogram above, decide whether the following statement is true or false:

>If you make 200 one-dollar bets on red, your chance of losing money is more than 50%.

Assign `loss_more_than_50` to either `True` or `False` depending on your answer to the question.

```python
loss_more_than_50 = ...
```

Tests for q1_1 (points: [2, 2])
```
>>> dollar_bet_on_red('red') == 1
True
```
```
>>> dollar_bet_on_red('green') == -1
True
```


Tests for q1_10 (points: [0, 0, 4])
```
>>> # Make sure histogram_statements is an array.
>>> import numpy as np
>>> type(histogram_statements) == np.ndarray
True
```
```
>>> # Make sure you have set histogram_statements to an array with at least 1 number
>>> 1 <= histogram_statements.item(0) <= 3
True
```
```
>>> histogram_statements.item(0) == 1 and histogram_statements.item(1) == 2 and histogram_statements.item(2) == 3
True
```


Tests for q1_2 (points: [0, 4])
```
>>> # Make sure your column names are correct
>>> wheel.labels[2] == "Winnings: Red"
True
```
```
>>> sum(wheel.column("Winnings: Red")) == -2
True
```


Tests for q1_3 (points: [0, 4])
```
>>> # If this test isn't passing, try running the cells from the top.
>>> set(["Pocket", "Color", "Winnings: Red"]) == set(ten_bets.labels)
True
```
```
>>> ten_bets.num_rows == 10
True
```


Tests for q1_4 (points: [0, 4])
```
>>> net_gain_red(10000) != net_gain_red(10000)
True
```
```
>>> -10000 <= net_gain_red(10000) <= 10000
True
```


Tests for q1_5 (points: [0, 4])
```
>>> len(all_gains_red) == 10000
True
```
```
>>> np.count_nonzero(all_gains_red <= 100) == 10000
True
```


Tests for q1_6 (points: [0, 4])
```
>>> type(loss_more_than_50) == bool
True
```
```
>>> loss_more_than_50 == True
True
```


Tests for q1_7 (points: [2, 1, 1])
```
>>> dollar_bet_on_split('5') == 17
True
```
```
>>> dollar_bet_on_split('6') == 17
True
```
```
>>> dollar_bet_on_split('10') == -1
True
```


Tests for q1_8 (points: [0, 4])
```
>>> # Make sure your column names are correct
>>> wheel.labels[3] == "Winnings: Split"
True
```
```
>>> sum(wheel.column("Winnings: Split")) == -2
True
```


Tests for q1_9 (points: [0, 4])
```
>>> len(all_gains_split) == 10000
True
```
```
>>> np.count_nonzero(all_gains_split >= -200) == 10000
True
```


Tests for q2_1 (points: [0, 4])
```
>>> # Your probability output should be a value between 0 and 1.
>>> 0 < first_three_black < 1
True
```
```
>>> first_three_black
0.10628371482723427
```


Tests for q2_2 (points: [0, 4])
```
>>> # Your probability output should be a value between 0 and 1.
>>> 0 < no_green < 1
True
```
```
>>> 0.582 < no_green < 0.583
True
```


Tests for q2_3 (points: [0, 4])
```
>>> # Your probability output should be a value between 0 and 1.
>>> 0 < at_least_one_green < 1
True
```
```
>>> 0.417643346770 <= at_least_one_green <= 0.417643346771
True
```


Tests for q2_4 (points: [0, 4])
```
>>> # Your probability output should be a value between 0 and 1.
>>> 0 < lone_winners < 1
True
```
```
>>> lone_winners
0.0011374153385028288
```


Tests for q2_5 (points: [0, 0, 4])
```
>>> # Make sure you are setting toys_option to an int
>>> type(toys_option) == int
True
```
```
>>> # toys_option should be assigned to either 1, 2 or 3.
>>> 1 <= toys_option <= 3
True
```
```
>>> toys_option == 1
True
```


Tests for q2_6 (points: [0, 0, 4])
```
>>> # Make sure you are setting lottery_option to an int
>>> type(lottery_option) == int
True
```
```
>>> # toys_option should be assigned to either 1, 2 or 3.
>>> 1 <= lottery_option <= 3
True
```
```
>>> lottery_option == 2
True
```


Tests for q2_7 (points: [0, 0, 4])
```
>>> # Make sure you are setting coin_option to an int
>>> type(coin_option) == int
True
```
```
>>> 1 <= coin_option <= 3
True
```
```
>>> coin_option == 1
True
```


Tests for q4_1 (points: [0, 0, 1, 3])
```
>>> # The array should have length 2
>>> len(deck_model_probabilities) == 2
True
```
```
>>> # The elements in the array should add up to 1.
>>> sum(deck_model_probabilities) == 1
True
```
```
>>> deck_model_probabilities.item(0) == 4/13
True
```
```
>>> deck_model_probabilities.item(1) == 9/13
True
```


Tests for q4_2 (points: [0, 0, 5])
```
>>> # Make sure you are setting statistic_choice to an int
>>> type(statistic_choice) == int
True
```
```
>>> # statistic_choice should be assigned to either 1, 2 or 3.
>>> 1 <= statistic_choice <= 3
True
```
```
>>> statistic_choice == 3
True
```


Tests for q4_3 (points: [5])
```
>>> # The statistic should be between 0 and 13 face cards for
>>> # a sample size of 13
>>> num_face = deck_simulation_and_statistic(13, deck_model_probabilities)
>>> 0 <= num_face <= 13
True
```


Tests for q4_4 (points: [2, 3])
```
>>> # There should be exactly as many elements in deck_statistics
>>> # as the number 'repetitions'
>>> len(deck_statistics) == repetitions
True
```
```
>>> # Each element of deck_statistics should be between 0
>>> # and 13 inclusive
>>> all([0 <= k <= 13 for k in deck_statistics])
True
```
